# Demo: How nbgrader Notebooks Work

Welcome to the course assignments! This short demo walks you through the three types of cells you will encounter, then gives you three quick practice exercises (15 points total) so you can verify that your setup works before tackling the real assignments.

## How this notebook works

Every assignment notebook contains three kinds of cells:

1. **Read-only cells** (locked, gray background). These load packages, define parameters, or set up the problem. Do not modify them.
2. **Solution cells**. These contain the placeholder `# YOUR CODE HERE` followed by `raise NotImplementedError()`. Delete both lines and replace them with your code. Keep the function signature and variable names exactly as given.
3. **Test cells** (locked). Some assertions are visible so you can check your work locally. Additional hidden tests run on the autograder server and are worth the bulk of the points.

**Before submitting**, always run **Kernel → Restart & Run All** to make sure every cell executes from a clean state.

## Setup

```python
import numpy as np
import matplotlib.pyplot as plt

plt.rcParams['figure.figsize'] = (8, 5)
plt.rcParams['font.size'] = 11
```

## Exercise 1: Compound interest (5 points)

A principal $P$ invested at annual interest rate $r$ for $t$ years grows to

$$A = P(1 + r)^{t}.$$

Implement a function that returns the future value $A$.

<!-- grade_id: ex_1, points: 0, grade: false, solution: true -->
```python
def compound_interest(P, r, t):
    """Return the future value of principal P at rate r after t years."""
    # YOUR CODE HERE
    raise NotImplementedError()


# Quick sanity check
print(f"$100 at 5% for 10 years = ${compound_interest(100, 0.05, 10):.2f}")
```

<!-- grade_id: test_1_visible, points: 3, grade: true, solution: false -->
```python
# Autograder tests for Exercise 1 (visible)
assert callable(compound_interest), "compound_interest must be a function"
assert np.isclose(compound_interest(100, 0.05, 0), 100), "With t=0 the principal should be unchanged"
assert np.isfinite(compound_interest(100, 0.05, 10)), "Result should be finite"
```

<!-- grade_id: test_1_hidden, points: 2, grade: true, solution: false -->
```python

```

## Exercise 2: Present discounted value (5 points)

The present discounted value of a stream of cash flows $C_0, C_1, \ldots, C_{T-1}$ at discount rate $r$ is

$$\text{PDV} = \sum_{t=0}^{T-1} \frac{C_t}{(1+r)^{t}}.$$

Implement a function that takes an array of cash flows and returns the PDV.

<!-- grade_id: ex_2, points: 0, grade: false, solution: true -->
```python
def pdv(cashflows, r):
    """Return the present discounted value of a cashflow array at rate r."""
    # YOUR CODE HERE
    raise NotImplementedError()


# Quick sanity check
print(f"PDV of [100, 100, 100] at r=0.05: ${pdv([100, 100, 100], 0.05):.2f}")
```

<!-- grade_id: test_2_visible, points: 2, grade: true, solution: false -->
```python
# Autograder tests for Exercise 2 (visible)
assert callable(pdv), "pdv must be a function"
assert np.isclose(pdv([100], 0.05), 100), "A single cash flow at t=0 should equal itself"
assert np.isfinite(pdv([100, 100, 100], 0.05)), "Result should be finite"
```

<!-- grade_id: test_2_hidden, points: 3, grade: true, solution: false -->
```python

```

## Exercise 3: Plot a budget constraint (5 points)

A consumer earns income $Y$ in period 1 (and zero in period 2) and can save or borrow at gross interest rate $R = 1 + r$. The budget constraint in period 2 is

$$c_2 = (Y - c_1) \cdot R.$$

Compute two arrays:
- `c1_vals`: 200 evenly spaced values from 0 to $Y$ (inclusive)
- `c2_vals`: the corresponding period-2 consumption from the budget constraint

Then plot the budget line.

<!-- grade_id: ex_3, points: 0, grade: false, solution: true -->
```python
# Parameters (do not change)
Y = 100     # period-1 income
R = 1.05    # gross interest rate

# YOUR CODE HERE
raise NotImplementedError()

# Plot the budget constraint
plt.plot(c1_vals, c2_vals)
plt.xlabel('Period 1 consumption ($c_1$)')
plt.ylabel('Period 2 consumption ($c_2$)')
plt.title('Intertemporal budget constraint')
plt.grid(True, alpha=0.3)
plt.show()
```

<!-- grade_id: test_3_visible, points: 2, grade: true, solution: false -->
```python
# Autograder tests for Exercise 3 (visible)
assert len(c1_vals) == len(c2_vals), "c1_vals and c2_vals must have the same length"
assert np.all(c1_vals >= 0), "c1_vals should be non-negative"
assert np.all(c2_vals >= -1e-10), "c2_vals should be non-negative along the budget line"
```

<!-- grade_id: test_3_hidden, points: 3, grade: true, solution: false -->
```python

```

## Submission

1. Re-run the entire notebook: **Kernel → Restart & Run All**
2. Verify that every cell executes without errors
3. Save the notebook
4. Upload the `.ipynb` file to Gradescope